<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>8.2. Вежбање &#8212; Petlja - Настава програмирања за Први разред гимназије</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="17. лекција - квиз" href="../testovi/index.html" />
    <link rel="prev" title="8.1. Основне аритметичке операције и примена" href="8a.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Настава програмирања за Први разред гимназије';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Настава програмирања за Први разред гимназије';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Настава програмирања за Први разред гимназије</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">1. лекција - Робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1a.html">1.1. Робот Карел - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1b.html">1.2. Робот Карел - бројачка петља <cite>for</cite></a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">2. лекција - Робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2a.html">2.1. Робот Карел - условне петље, гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2b.html">2.2. Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2b.html#id11">2.3. Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">3. лекција - вежбање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3a.html">3.1. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">4. Квиз - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4a.html">4.1. Квиз - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">5. лекција - Корњача графика (гранање, угнежљене петље, процедуре, торке, листе)</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5a.html">5.1. Корњача графика - гранање, угнежђене петље, процедуре, торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html">5.2. Процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html#id9">5.3. Торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html#id15">5.4. Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../6_lekcija/index.html">6. лекција - вежбање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6a.html">6.1. Корњача графика - вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">7. лекција -квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7a.html">7.1. Квиз - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">8. лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="8a.html">8.1. Основне аритметичке операције и примена</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">8.2. Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="#id20">Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">9. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../9_lekcija/index.html">9. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9a.html">9.1. Цели и реални бројеви, дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9b.html">9.2. Целобројно дељење - позициони запис броја</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../10_lekcija/index.html">10. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10a.html">10.1. Коришћење готових функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10b.html">10.2. Дефинисање функција</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">11. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">12. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12a.html">12.1. Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12b.html">12.2. Конструкција <cite>elif</cite></a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../13_lekcija/index.html">13. лекција - понављање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13a.html">13.1. Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13b.html">13.2. Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13b.html#id7">13.3. Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">14. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">15. лекција - структуре</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15a.html">15.1. Торке, листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15a.html#id49">15.2. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15b.html">15.3. Скупови, речници</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../16_lekcija/index.html">16. лекција - структуре</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../16_lekcija/16a.html">16.1. Обрада структура података коришћењем петљи</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">17. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>8.2. Вежбање<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Технике које смо до сада научили довољне су нам да бисмо решили велики број
математичких задатака. Размотримо неколико.</p>
<div class="section" id="id2">
<h2>Године маме и тате<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Милица има 4 године, њена мама има 7 пута више година него она, а њен
тата има 8 пута више година него она. Колико је година Миличин тата старији
од њене маме?</p>

    </p></div></div>

        <div class="course-box course-box-problem course-content" style="max-width: none;">
        <pre data-component="parsons" id="godine"      >
        <span data-question> Q-37: Поређај делове кода тако да представљају исправно решење овог задатка.</span>milica = 4
---
mama = 7 * milica
tata = 8 * milica
---
razlika = tata - mama
---
print(razlika)
        </pre>
        </div>
    </div>
<div class="section" id="id3">
<h2>Река Морава<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Велика Морава је дугачка 185km и настаје од Јужне Мораве, која је
90km дужа, и Западне Мораве, која је 123km дужа од ње. Колика је
укупна дужина ове три реке?</p>

    </p></div></div>

<div data-childcomponent="морава" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=морава data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
velika_morava = 185
juzna_morava = velika_morava + 90
zapadna_morava = velika_morava + 123
ukupno = ??? # ispravi ovaj red
print(ukupno)
</textarea>
</div>
</div>
<div class="section" id="id5">
<h2>Немањићи<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Стефан Немањић је постао краљ Србије 1217 и владао је 11
година. После њега је Радослав владао до 1234. године, па Владислав,
који је владао 9 година и предао престо брату Урошу Првом, који је
владао до 1276. У којим временским периодима су владали ови српски
краљеви?</p>

    </p></div></div>

<div data-childcomponent="немањићи" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=немањићи data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
Stefan_pocetak = 1217
Stefan_kraj = 1217 + 11
Radoslav_pocetak = Stefan_kraj
Radoslav_kraj = 1234
Vladislav_pocetak = 0
Vladislav_kraj = 0
Uros_pocetak = 0
Uros_kraj = 0
print("Стефан:", Stefan_pocetak, "-", Stefan_kraj)
print("Радослав:", Radoslav_pocetak, "-", Radoslav_kraj)
print("Владислав:", Vladislav_pocetak, "-", Vladislav_kraj)
print("Урош:", Uros_pocetak, "-", Uros_kraj)
</textarea>
</div>
<p>Исправи претходни програм тако да исправно израчуна периоде у којима
су владали краљеви. Ако све урадиш како треба добићеш следеће резултате:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Стефан</span><span class="p">:</span> <span class="mi">1217</span> <span class="o">-</span> <span class="mi">1228</span>
<span class="n">Радослав</span><span class="p">:</span> <span class="mi">1228</span> <span class="o">-</span> <span class="mi">1234</span>
<span class="n">Владислав</span><span class="p">:</span> <span class="mi">1234</span> <span class="o">-</span> <span class="mi">1243</span>
<span class="n">Урош</span><span class="p">:</span> <span class="mi">1243</span> <span class="o">-</span> <span class="mi">1276</span>
</pre></div>
</div>
</div>
<div class="section" id="print">
<h2>Напредније коришћење <code class="docutils literal notranslate"><span class="pre">print</span></code><a class="headerlink" href="#print" title="Permalink to this headline">¶</a></h2>

    <div data-level="3" style="display:none">

    </div>
<p>Приметимо да смо у претходном задатку додали испис имена краљева и
цртица између почетка и краја њихове владавине тако што смо тај текст
који смо желели да се испише ставили под знаке навода (нпр. навели смо
<code class="docutils literal notranslate"><span class="pre">&quot;Стефан:</span> <span class="pre">&quot;</span></code>). О раду са текстом ће више речи бити касније. Такође,
приметимо да смо овај пут навели неколико ствари унутар <code class="docutils literal notranslate"><span class="pre">print</span></code>,
раздвојених зарезима. У тим ситуацијама <code class="docutils literal notranslate"><span class="pre">print</span></code> штампа сваку од њих,
развајајући их размацима (на пример, када се изврши нареба
<code class="docutils literal notranslate"><span class="pre">print(&quot;baci&quot;,</span> <span class="pre">5)</span></code> исписује се <code class="docutils literal notranslate"><span class="pre">baci</span> <span class="pre">5</span></code>).</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Као што смо рекли, ствари наведене унутар <code class="docutils literal notranslate"><span class="pre">print</span></code> раздвајају
се са по једним размаком. То се може променити тако што се на
крају <code class="docutils literal notranslate"><span class="pre">print</span></code> наведе <code class="docutils literal notranslate"><span class="pre">sep=&quot;&quot;</span></code> и унутар наводника наведе шта
ће се користити да раздвоји делове. На пример, ако се наведе
<code class="docutils literal notranslate"><span class="pre">print(1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">sep=&quot;&quot;)</span></code> исписаће се <code class="docutils literal notranslate"><span class="pre">123</span></code>, а ако се наведе
<code class="docutils literal notranslate"><span class="pre">print(1,</span> <span class="pre">2,</span> <span class="pre">3,</span> <span class="pre">sep=&quot;,</span> <span class="pre">&quot;)</span></code> исписаће се <code class="docutils literal notranslate"><span class="pre">1,</span> <span class="pre">2,</span> <span class="pre">3</span></code>. Након
сваког извршавања <code class="docutils literal notranslate"><span class="pre">print</span></code>, прелази се у нови ред (наредни
позиви <code class="docutils literal notranslate"><span class="pre">print</span></code> штампаће свој резултат у наредном реду). И то
се може променити тако што се на крају <code class="docutils literal notranslate"><span class="pre">print</span></code> наведе
<code class="docutils literal notranslate"><span class="pre">end=&quot;&quot;</span></code> и унутар наводника оно што ће се користити након
целог исписа. На пример, <code class="docutils literal notranslate"><span class="pre">print(1,</span> <span class="pre">2,</span> <span class="pre">end=&quot;&quot;)</span></code> проузрокује да
се након исписа не пређе у нови ред, већ да наредни испис иде
непосредно након вредности <code class="docutils literal notranslate"><span class="pre">2</span></code>.</p>

    </p></div></div>
</div>
<div class="section" id="id7">
<h2>Једначина<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Напиши програм који израчунава који број треба додати броју 123780
да се добије број 321732.</p>

    </p></div></div>
<p>Нажалост, Python не може директно да решава једначине. Ти мораш да
напишеш израз којим се непозната вредност израчунава на основу
познатих, а онда ти он може помоћи у рачунању.</p>

<div data-childcomponent="непознати_сабирак" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=непознати_сабирак data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
prvi_sabirak = 123780
zbir = 321732
drugi_sabirak = 0    # popravi resenje
print(drugi_sabirak)
</textarea>
</div>
<p>Провери да ли је твој програм израчунао тачно решење.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_једначина">
            <p>Колико је решење?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^197952$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041e\u0434 \u0437\u0431\u0438\u0440\u0430 \u043e\u0434\u0443\u0437\u043c\u0438 \u043f\u043e\u0437\u043d\u0430\u0442\u0438 \u0441\u0430\u0431\u0438\u0440\u0430\u043a\u201d</p>\n"}]]
        </script>

        </div>
        </div>
            <p>Ако у решењу нису коришћене вредности, већ само називи променљивих,
програм би требало да исправно решава задатке и за друге
бројеве. Тестирај га на тест-примерима које смо припремили.</p>

<div data-childcomponent="непознати_сабирак_тест" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=непознати_сабирак_тест data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="prvi_sabirak zbir drugi_sabirak"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
prvi_sabirak = 123780
zbir = 321732
# -*- acsection: main -*-
drugi_sabirak = 0    # popravi resenje
# -*- acsection: after-main -*-
print(drugi_sabirak)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for prvi_sabirak, zbir, drugi_sabirak in [(100, 230, 130), (200, 942, 742)]:
          self.assertEqual(acMainSection(prvi_sabirak = prvi_sabirak, zbir = zbir)["drugi_sabirak"],drugi_sabirak,"Ако је једначина %s + x = %s, тада је x = %s." % (prvi_sabirak, zbir, drugi_sabirak))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id10">
<h2>Фудбалски терен<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>
<p>Илуструјмо кроз решење наредног, мало сложенијег задатка, разне
приступе решавању једног те истог задатка и продискутујмо предности и
мане разних решења.</p>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Дужина фудбалског терена је 115 метара, а ширина 80 метара. Ана трчи
по правоугаоној стази која је са сваке стране терена споља удаљена
по 5 метара. Колико она претрчи, ако се зна да је оптрчала терен 3
пута.</p>

    </p></div></div>
<a class="reference internal image-reference" href="../_images/teren.jpg"><img alt="../_images/teren.jpg" class="align-center" src="../_images/teren.jpg" style="width: 400px;" /></a>
<p>Ана претрчи троструку дужину обима стазе. Да бисмо израчунали обим
стазе потребно је да израчунамо њену дужину и ширину. Стаза је 10
метара дужа од терена (по 5 метара са сваке стране), и 10 метара шира
од терена (по 5 метара са сваке стране). Обим је једнак двострукој
вредности збира дужине и ширине, па задатак можемо решити на следећи
начин.</p>

<div data-childcomponent="терен_1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(3 * 2 * (125 + 90))
</textarea>
</div>
<p>Овај програм коректно израчунава претрчан пут, међутим, има неколико
недостатака.  Прво, неке величине смо рачунали у глави, уместо да
препустимо рачунару да их израчуна (на пример, 125 смо добили као
115 + 2·5), чиме смо себи створили непотребан посао и, што је још
важније, увели у програм могуће рачунске грешке (за разлику од
рачунара који у рачуну не греше баш никада, људи у рачуну често
греше). Даље, читајући текст програма није баш јасно зашто се резултат
добија баш на начин на који је то урађено у програму. Такође, овакав
програм није потпуно једноставно изменити. Размислимо шта би се десило
када би се улазни подаци мало променили, тј. када би бисмо посматрали
терен дугачак 110 метара, широк 80 метара и стазу која је на растојању
10 метара од терена. Да бисмо изменили програм, потребно би било да
поново у глави израчунамо дужину и ширину стазе, што је непотребно
компликовано.</p>
<p>Наредни програм је мало бољи од претходног, али и даље је неразумљив,
а да бисмо га променили морамо бити пажљиви и улазне податке изменити
на свим местима (на пример, број 5 је потребно заменити два пута).</p>

<div data-childcomponent="терен_2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(3 * 2 * ((110 + 2 * 5) + (80 + 2 * 5)))
</textarea>
</div>
<p>Погледајмо сада другачије решење истог задатка.</p>

<div data-childcomponent="терен_3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
duzina_terena = 115
sirina_terena = 80
rastojanje = 5
broj_krugova = 3
duzina_staze = duzina_terena + 2 * rastojanje
sirina_staze = sirina_terena + 2 * rastojanje
obim_staze = 2 * (duzina_staze + sirina_staze)
pretrcan_put = broj_krugova * obim_staze
print(pretrcan_put)
</textarea>
</div>
<p>Приметимо неколико ствари. Програм је дужи, па га је сигурно било теже
написати. Све улазне податке (дужину и ширину терена, као и
растојање стазе од терена) смо именовали на самом почетку
програма. Такође, именовали смо и све међурезултате (дужину и ширину
стазе, њен обим и претрчан пут), чиме смо постигли да је оном ко чита
програм сада потпуно јасно како се заправо дошло до траженог резултата
и зашто су формуле такве какве јесу. Помоћу дугмета <strong>Корак по корак</strong>
лако можемо сазнати све међурезултате које смо именовали, а ако у
неком тренутку то буде потребно и њих можемо једноставно исписати
коришћењем <code class="docutils literal notranslate"><span class="pre">print</span></code>. Такође, постигли смо и то да је измене улазних
података сада потпуно једноставно урадити - само је потребно променити
податке дате на почетку програма.</p>
<p>Као што смо видели, наш систем ти омогућава да се твој програм
аутоматски тестира за различите улазне податке. Покушај сада да
допуниш наредни програм и да покренеш механизам тестирања и видећеш да
ли твој програм ради исправно за различите улазне податке.</p>

<div data-childcomponent="терен_3_test" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_3_test data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="duzina_terena sirina_terena rastojanje broj_krugova pretrcan_put"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
duzina_terena = 115
sirina_terena = 80
rastojanje = 5
broj_krugova = 3
# -*- acsection: main -*-
duzina_staze = 0  # popravi ovaj red
sirina_staze = 0  # popravi ovaj red
obim_staze = 0    # popravi ovaj red
pretrcan_put = 0  # popravi ovaj red
# -*- acsection: after-main -*-
print(pretrcan_put)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for duzina_terena, sirina_terena, rastojanje, broj_krugova, pretrcan_put in [(110, 80, 7, 4, 1744), (75, 65, 3, 10, 3040)]:
          self.assertEqual(acMainSection(duzina_terena = duzina_terena, sirina_terena = sirina_terena, rastojanje = rastojanje, broj_krugova = broj_krugova)["pretrcan_put"],pretrcan_put,
          "Ако је терен дугачак %s, а широк %s метара и ако Ана трчи %s круга на %s метара од ивице, тада она пређе %s метара." % (duzina_terena, sirina_terena, broj_krugova, rastojanje, pretrcan_put))
myTests().main()
</textarea>
</div>
<p>На крају, можемо направити програм који ће при свом покретању питати
корисника да унесе одговарајуће податке.</p>

<div data-childcomponent="терен_4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
duzina_terena = int(input("Unesi dužinu terena:"))
sirina_terena = int(input("Unesi širinu terena:"))
rastojanje    = int(input("Unesi rastojanje od staze do terena:"))
broj_krugova  = int(input("Unesi broj krugova koje Ana pretrči:"))
duzina_staze = duzina_terena + 2 * rastojanje
sirina_staze = sirina_terena + 2 * rastojanje
obim_staze = 2 * (duzina_staze + sirina_staze)
pretrcan_put = broj_krugova * obim_staze
print(pretrcan_put)
</textarea>
</div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Покушај сада да измениш претходни програм тако што ћеш
претпоставити да Ана трчи по правоугаоној стази удаљеној неколико
метара од ивице, али на унутра. Претпостави да су дужина и ширина
терена (<code class="docutils literal notranslate"><span class="pre">duzina_terena</span></code>, <code class="docutils literal notranslate"><span class="pre">sirina_terena</span></code>), растојање стазе од
ивице терена (<code class="docutils literal notranslate"><span class="pre">rastojanje</span></code>), као и број кругова које треба да
претрчи (<code class="docutils literal notranslate"><span class="pre">broj_krugova</span></code>) задати на почетку текста
програма. Измени само редове који су обележени, а остало не
дирај. Ако све урадиш како треба, након покретања програма
(дугметом <strong>Покрени програм</strong>) и аутоматског тестирања сва поља ће
бити зелене боје.</p>

    </p></div></div>

<div data-childcomponent="терен_унутра_тест" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_унутра_тест data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="duzina_terena sirina_terena rastojanje broj_krugova pretrcan_put"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
duzina_terena = 115
sirina_terena = 80
rastojanje = 5
broj_krugova = 3
# -*- acsection: main -*-
duzina_staze = 0   # popravi ovaj red
sirina_staze = 0   # popravi ovaj red
obim_staze = 0     # popravi ovaj red
pretrcan_put = 0   # popravi ovaj red
# -*- acsection: after-main -*-
print(pretrcan_put)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for duzina_terena, sirina_terena, rastojanje, broj_krugova, pretrcan_put in [(100, 80, 5, 2, 640), (95, 85, 20, 7, 1400), (110, 70, 12, 4, 1056)]:
          self.assertEqual(acMainSection(duzina_terena = duzina_terena, sirina_terena = sirina_terena, rastojanje = rastojanje, broj_krugova = broj_krugova)["pretrcan_put"],pretrcan_put,
          "Ако је терен дугачак %s, а широк %s метара и ако Ана трчи %s круга на %s метара од ивице, тада она пређе %s метара." % (duzina_terena, sirina_terena, broj_krugova, rastojanje, pretrcan_put))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="терен_тест_решење"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    <p>Ево једног начина да се дође до тачног решења.</p>

<div data-childcomponent="терен_тест_решење1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=терен_тест_решење1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
duzina_terena = 100
sirina_terena = 80
rastojanje = 5
broj_krugova = 2

duzina_staze = duzina_terena - 2*rastojanje
sirina_staze = sirina_terena - 2*rastojanje
obim_staze = 2 * (sirina_staze + duzina_staze)
pretrcan_put = broj_krugova * obim_staze

print(pretrcan_put)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id17">
<h2>Позиви на концерт<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Мирјана и Лидија певају у истом хору и припремају се за концерт.
Мирјана има 245 пријатеља на једној друштвеној мрежи, док их Лидија
има 218. Када је Мирјана погледала Лидијин профил, видела је да
имају 114 заједничких пријатеља. Ако би и једна и друга позвале све
своје пријатеље са те друштвене мреже на концерт, колико различтих
људи би добило тај позив.</p>

    </p></div></div>
<p>Скуп људи који ће добити позив на концерт можемо поделити на скуп
Мирјаниних пријатеља који нису уједно и Лидијини, на скуп Лидијиних
пријатеља који нису уједно и Мирјанини и на скуп њихових заједничких
пријатеља (ова три скупа су дисјунктна тј. не постоји ни један
пријатељ који је члан више ових скупова). На основу датих података
можемо израчунати колико Мирјана има пријатеља који нису уједно и
Лидијини пријатељи (као разлику укупног броја Мирјаниних пријатеља и
броја њихових заједничких пријатеља). Слично можемо израчунати и број
Лидијини пријатеља који нису уједно и Мирјанини пријатељи. Пошто су
три описана скупа дисјунктна, yкупан број људи који ће бити позвани на
концерт је збир ова два броја и броја заједничких пријатеља који нам
је од почетка познат.</p>

<div data-childcomponent="заједнички_пријатељи" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=заједнички_пријатељи data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="mirjanini lidijini zajednicki ukupno"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
mirjanini = 245
lidijini = 218
zajednicki = 114
# -*- acsection: main -*-
samo_mirjanini = mirjanini - zajednicki
samo_lidijini = lidijini - zajednicki
ukupno = 0 # ispravi ovaj red
# -*- acsection: after-main -*-
print(ukupno)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for mirjanini, lidijini, zajednicki, ukupno in [(200, 200, 100, 300), (150, 130, 80, 200), (73, 42, 15, 100)]:
          self.assertEqual(acMainSection(mirjanini=mirjanini, lidijini=lidijini, zajednicki=zajednicki)["ukupno"],ukupno,
          "Ако је Мирјана позвала %s пријатеља, Лидија %s, а заједничких је %s, тада је укупно позвано %s људи." % (mirjanini, lidijini, zajednicki, ukupno))
myTests().main()
</textarea>
</div>
<p>Приметимо да је задатак могуће решити и једноставније. Наиме скуп људи
који ће добити позив можемо разложити на скуп свих Мирјаниних
пријатеља и скуп свих Лидијиних пријатеља који нису уједно Мирјанини.
Ова два скупа су дисјунктна, тако да је резултат могуће добити као
разлику између збира Мирјаниних и Лидијиних пријатеља и броја њихових
заједничких пријатеља.</p>

<div data-childcomponent="заједнички_пријатељи_1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=заједнички_пријатељи_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
mirjanini = 245
lidijini = 218
zajednicki = 114
ukupno = mirjanini + lidijini - zajednicki
print(ukupno)
</textarea>
</div>
<p>Формула коју смо применили у овом задатку назива се <em>формула
укључења-искључења</em> и она каже да за било која два скупа A и B важи да
је |A ∪ B| = |A| + |B| - |A ∩ B|.</p>
</div>
</div>
<div class="section" id="id20">
<h1>Домаћи задатак<a class="headerlink" href="#id20" title="Permalink to this headline">¶</a></h1>
<p>Ова лекција није много тешка, али је прилично обимна, па је сасвим
могуће да неки од приказаних задатака нисте стигли да проанализирате
током часа. Ако има таквих задатака, уради их сада, у склопу домаћег
задатка. Након тога уради и наредне задатке. Додатне задаткеза вежбу
можеш пронаћи <a class="reference external" href="IzracunavanjeZadaci.html">овде</a>.</p>
<div class="section" id="id22">
<h2>Израз<a class="headerlink" href="#id22" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Збир бројева 23765 и 7825 умањи 45 пута, па добијени број повећај
за 1609. Колики је резултат?  Задатак реши једним изразом (немој да
рачунаш пешке).</p>

    </p></div></div>

<div data-childcomponent="израз_2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=израз_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print() # у заграде упиши израз
</textarea>
</div>
<p>Провери да ли је твој програм израчунао тачно решење.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_израз2">
            <p>Колико је решење?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^2311$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="id24">
<h2>Тркачи<a class="headerlink" href="#id24" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Васа је прешао 2347 метара. Воја 987 метара више од Васе, а Милош два
пута више од Воје. Колико су метара укупно прешли?</p>

    </p></div></div>

<div data-childcomponent="три_тркача" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=три_тркача data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="vasa ukupno"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
vasa  = 2347
# -*- acsection: main -*-
# dopuni ovde kod
# -*- acsection: after-main -*-
print(ukupno)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for vasa, ukupno in [(2462, 12809), (773, 6053)]:
          self.assertEqual(acMainSection(vasa = vasa)["ukupno"],ukupno,"Ако је Васа претрчао %s метара, укупно су претрчали %s метара." % (vasa, ukupno))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="тркачи_решење_reveal"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="три_тркача_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=три_тркача_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
vasa  = 2347
voja  = vasa + 987
milos = voja * 2
ukupno = vasa + voja + milos
print(ukupno)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id27">
<h2>Тепих<a class="headerlink" href="#id27" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У средини собе квадратног облика се налази тепих. Ако је позната димензија
собе у метрима, удаљеност тепиха од зидова, и цена прања по једном квадратном
метру тепиха, израчунај цену прања целог тепиха.</p>

    </p></div></div>

<div data-childcomponent="тепих" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=тепих data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="dimenzija_sobe udaljenost_tepiha cena_pranja_po_m2 ukupna_cena_pranja"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
dimenzija_sobe = int(input("Unesi dimenziju sobe:"))
udaljenost_tepiha = int(input("Unesi udaljenost tepiha od zida:"))
cena_pranja_po_m2 = int(input("Unesi cenu pranja kvadratnog metra tepiha:"))
# -*- acsection: main -*-
# zavrsi program
# -*- acsection: after-main -*-
print(ukupna_cena_pranja)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for dimenzija_sobe, udaljenost_tepiha, cena_pranja_po_m2, ukupna_cena_pranja in [(6, 2, 120, 480), (4, 1, 150, 600)]:
          self.assertEqual(acMainSection(dimenzija_sobe = dimenzija_sobe, udaljenost_tepiha = udaljenost_tepiha, cena_pranja_po_m2 = cena_pranja_po_m2)["ukupna_cena_pranja"],ukupna_cena_pranja,
          "Ако је димензија собе %s метара, удаљеност тепиха од зидова %s метара, цена прања по метру квадратном %s динара, тада је укупна цена прања %s динара." % (dimenzija_sobe, udaljenost_tepiha, cena_pranja_po_m2, ukupna_cena_pranja))
myTests().main()
</textarea>
</div>
<p>Када га покренеш и унесеш бројеве 5, 1 и 15 треба да добијеш
резултат 135. Пробај и да га аутоматски тестираш на основу
тест-примера које смо ти припремили.</p>
</div>
<div class="section" id="id29">
<h2>Фудбалери и кошаркаши<a class="headerlink" href="#id29" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У одељењу има 15 дечака и сви они тренирају или фудбал или
кошарку. Ако се зна да фудбал тренира 10 ученика, а кошарку тренира
8 ученика колико је дечака који тренирају оба спорта?</p>

    </p></div></div>
<p>Погледајмо наредну табелу која нам може помоћи да нађемо правило како
се може одредити број дечака који тренирају два спорта (са ф смо
обележили фудбалере, а са к кошаркаше).</p>
<table class="docutils align-default">
<colgroup>
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
<col style="width: 7%" />
</colgroup>
<tbody>
<tr class="row-odd"><td><p>1</p></td>
<td><p>2</p></td>
<td><p>3</p></td>
<td><p>4</p></td>
<td><p>5</p></td>
<td><p>6</p></td>
<td><p>7</p></td>
<td><p>8</p></td>
<td><p>9</p></td>
<td><p>10</p></td>
<td><p>11</p></td>
<td><p>12</p></td>
<td><p>13</p></td>
<td><p>14</p></td>
<td><p>15</p></td>
</tr>
<tr class="row-even"><td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td><p>ф</p></td>
<td></td>
<td></td>
<td></td>
<td></td>
<td></td>
</tr>
<tr class="row-odd"><td></td>
<td></td>
<td></td>
<td></td>
<td></td>
<td></td>
<td></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
<td><p>к</p></td>
</tr>
</tbody>
</table>
<p>Уместо нуле упиши израз којим можемо да израчунамо број дечака који
тренирају два спорта. Када се програм покрене тестираће се на три
тест-примера.  Тест-примери на којима програм даје тачан резултат биће
обележени зеленом бојом, а они на којима се добија нетачан резултат
биће обележени црвеном бојом. Ако се сва три примера обележе зелено,
значи да је веома вероватно да је твоје решење тачно.</p>

<div data-childcomponent="пресек_спортиста" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=пресек_спортиста data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="fudbal kosarka ukupno_decaka dva_sporta"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
ukupno_decaka = int(input("Unesi ukupan broj dečaka: "))
fudbal = int(input("Unesi broj dečaka koji treniraju fudbal: "))
kosarka = int(input("Unesi broj dečaka koji treniraju kosarku: "))
# -*- acsection: main -*-
# dovrši program izračunavajući broj dečaka koji treniraju dva sporta
# -*- acsection: after-main -*-
print(dva_sporta)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for (ukupno_decaka, fudbal, kosarka, dva_sporta) in [(15, 10, 8, 3), (20, 14, 14, 8), (20, 9, 11, 0)]:
          self.assertEqual(acMainSection(ukupno_decaka=ukupno_decaka, fudbal=fudbal, kosarka=kosarka)["dva_sporta"], dva_sporta,
          "Ако у одељењу од %s ученика има %s фудбалера и %s кошаркаша, онда %s ученика тренирају два спорта." % (ukupno_decaka, fudbal, kosarka, dva_sporta))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="пресек_решење"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    <p>Приметимо да се у решењу овог задатка заправо поново може применити
формула укључења-искључења коју смо извели у претходном. Овај пут
имамо познат број елемената сваког од два скупа и број елемената
њихове уније и на основу тога израчунавамо број елемената њиховог
пресека. Стога је једно могуће решење.</p>
<p><code class="docutils literal notranslate"><span class="pre">dva_sporta</span> <span class="pre">=</span> <span class="pre">(fudbal</span> <span class="pre">+</span> <span class="pre">kosarka)</span> <span class="pre">-</span> <span class="pre">ukupno_decaka</span></code></p>

    </div>
    </div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - 8.1. Основне аритметичке операције и примена' data-toggle="tooltip"><a href="8a.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - 17. лекција - квиз' data-toggle="tooltip"><a href="../testovi/index.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - 8.1. Основне аритметичке операције и примена' data-toggle="tooltip"><a href="8a.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - 17. лекција - квиз' data-toggle="tooltip"><a href="../testovi/index.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>