<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>9.2. Целобројно дељење - позициони запис броја &#8212; Petlja - Настава програмирања за Први разред гимназије</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="10. лекција" href="../10_lekcija/index.html" />
    <link rel="prev" title="9.1. Цели и реални бројеви, дељење" href="9a.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Настава програмирања за Први разред гимназије';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Настава програмирања за Први разред гимназије';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Настава програмирања за Први разред гимназије</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">1. лекција - Робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1a.html">1.1. Робот Карел - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1b.html">1.2. Робот Карел - бројачка петља <cite>for</cite></a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">2. лекција - Робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2a.html">2.1. Робот Карел - условне петље, гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2b.html">2.2. Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2b.html#id11">2.3. Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">3. лекција - вежбање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3a.html">3.1. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">4. Квиз - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4a.html">4.1. Квиз - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">5. лекција - Корњача графика (гранање, угнежљене петље, процедуре, торке, листе)</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5a.html">5.1. Корњача графика - гранање, угнежђене петље, процедуре, торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html">5.2. Процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html#id9">5.3. Торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5b.html#id15">5.4. Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../6_lekcija/index.html">6. лекција - вежбање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6a.html">6.1. Корњача графика - вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">7. лекција -квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7a.html">7.1. Квиз - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../8_lekcija/index.html">8. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8a.html">8.1. Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8b.html">8.2. Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8b.html#id20">Домаћи задатак</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">9. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">9. лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="9a.html">9.1. Цели и реални бројеви, дељење</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">9.2. Целобројно дељење - позициони запис броја</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../10_lekcija/index.html">10. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10a.html">10.1. Коришћење готових функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10b.html">10.2. Дефинисање функција</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">11. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">12. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12a.html">12.1. Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12b.html">12.2. Конструкција <cite>elif</cite></a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../13_lekcija/index.html">13. лекција - понављање</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13a.html">13.1. Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13b.html">13.2. Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13b.html#id7">13.3. Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">14. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">15. лекција - структуре</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15a.html">15.1. Торке, листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15a.html#id49">15.2. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15b.html">15.3. Скупови, речници</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../16_lekcija/index.html">16. лекција - структуре</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../16_lekcija/16a.html">16.1. Обрада структура података коришћењем петљи</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../testovi/index.html">17. лекција - квиз</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../testovi/strukture.html">17.1. Квиз - структуре података и алгоритми</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>9.2. Целобројно дељење - позициони запис броја<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Ако је укупан број центиметара био 123, тада је број метара 1, број
дециметара 2 и број центиметара 3. Решавањем претходног задатка смо
заправо одређивали појединачне цифре коришћене у запису тог
троцифреног броја. Приказана техника може бити уопштена тако да се
одређују све цифре и у запису дужих бројева.</p>
<div class="section" id="id2">
<h2>Цифре броја<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Чест начин откривања грешака при слању података је да се уз податке
које треба послати, пошаљу и одређени контролни подаци, израчунати
на основу самих података. Када прималац прими податке, он на основу
примљених података поново израчунава контролне податке и упоређује
их са контролним подацима које је примио. Ако се приликом преноса
података, услед неких сметњи, подаци случајно измене, прималац ће
то приметити тако што ће видети да се контролни подаци које је он
израчунао неће поклопити са онима које је примио. Сви подаци се у
рачунарима представљају помоћу бројева, а још од најранијег доба
рачунарства као метода контроле коришћен је збир цифара. На пример,
ако би податак који се шаље био број 12345, онда би се уз њега слао
и контролни број 15 (збир 1 + 2 + 3 + 4 + 5). Ако би приликом слања
нека цифра била случајно промењена (на пример, ако би прималац
грешком примио број 12335) тада би се и контролни број
највероватније разликовао (контролни број би у нашем примеру био
14). Напиши програм који за дати петоцифрени број одређује његов
контролни број (збир његових цифара).</p>

    </p></div></div>

<div data-childcomponent="цифре_броја" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=цифре_броја data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="broj kontrolni_broj"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
broj = 12345
# -*- acsection: main -*-
c0 = (broj // 1) % 10
c1 = (broj // 10) % 10
c2 = 0 # ispravi ovaj red
c3 = (broj // 1000) % 10
c4 = 0 # ispravi ovaj red
kontrolni_broj = c0 + c1 + c2 + c3 + c4
# -*- acsection: after-main -*-
print(kontrolni_broj)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for broj, kontrolni_broj in [(71425, 19), (33214, 13), (62040, 12)]:
          self.assertEqual(acMainSection(broj = broj)["kontrolni_broj"],kontrolni_broj,"За број %s контролни број је %s." % (broj, kontrolni_broj))
myTests().main()
</textarea>
</div>
<p>Приметимо да смо цифре броја одређивали веома слично као у претходном
задатку.  Цифре одређујемо сдесна налево, тако што делимо број са
тежином цифре (за цифру јединица број делимо са 1, десетица са 10,
стотина са 100 итд.) и проналазимо остатак при дељењу са 10.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Људи су контролу података примењивали и ручно, а сада рачунари
обављају такве контроле у овиру прецизно задатих поступака –
<em>протокола за размену података</em>. Иако је контролни збир
најједноставнији начин контроле, он не може да открије грешке до
којих може доћи услед случајне размене редоследа цифара (на пример,
ако се уместо броја 12345 грешком пошаље број 12435, контролни збир
оба броја ће бити исти и грешка неће бити примећена). Зато се
уместо збира понекад користе изрази облика <span class="math notranslate nohighlight">\(c_0 + 2c_1 +
3c_2 + 4c_3 + 5c_4\)</span>. Покушај да измениш претходни програм тако да
израчунава контролни број на овај начин.</p>

    </p></div></div>
</div>
<div class="section" id="id4">
<h2>Целобројно дељење - време и углови<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h2>
<p>За разлику од бројева и јединица мере које записујемо у систему чија
је основа 10, при раду са временом и угловима користимо систем чија је
основа број 60. Тако један сат има 60 минута, а један минут 60
секунди. Слично један степен има 60 угаоних минута, а један угаони
минут има 60 угаоних секунди. Прикажимо сада кроз неколико задатака
како можемо у програмима вршити израчунавања у којима учествују време
и углови.</p>
</div>
<div class="section" id="id5">
<h2>Конверзија сати и минута у минуте и обратно<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Ако се зна колико је тренутно сати и минута, израчунај колико је
минута протекло од претходне поноћи.</p>

    </p></div></div>
<p>Пошто у једном сату има 60 минута, довољно је да помоножиш број сати
са 60 и на то да додаш број минута.</p>

<div data-childcomponent="сати_и_минути_у_минуте" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=сати_и_минути_у_минуте data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="sati minuta minuta_od_ponoci"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
sati = 2
minuta = 60
# -*- acsection: main -*-
minuta_od_ponoci = 0 # ispravi ovaj red
# -*- acsection: after-main -*-
print(minuta_od_ponoci)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for sati, minuta, minuta_od_ponoci in [(14, 19, 859), (11, 13, 673), (23, 59, 1439)]:
          self.assertEqual(acMainSection(sati = sati, minuta = minuta)["minuta_od_ponoci"],minuta_od_ponoci,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
myTests().main()
</textarea>
</div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Ако се зна колико је минута протекло од претходне поноћи, израчунај
колико је тренутно сати и минута.</p>

    </p></div></div>
<p>Ако са <span class="math notranslate nohighlight">\(s\)</span> обележимо тренутни број сати, са <span class="math notranslate nohighlight">\(m\)</span> тренутни
број минута, а са <span class="math notranslate nohighlight">\(M\)</span> број минута протеклих од поноћи, тада важи
да је <span class="math notranslate nohighlight">\(M = s \cdot 60 + m\)</span>, при чему за <span class="math notranslate nohighlight">\(m\)</span> важи да је
број између <span class="math notranslate nohighlight">\(0\)</span> и <span class="math notranslate nohighlight">\(59\)</span>, што јасно указује на то да се
тражене вредности могу израчунати применом целобројног дељења.</p>

<div data-childcomponent="минути_у_сате_и_минуте" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=минути_у_сате_и_минуте data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="minuta_od_ponoci sati minuta"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
minuta_od_ponoci = 125
# -*- acsection: main -*-
sati = 0     # ispravi ovaj red
minuta = 0   # ispravi ovaj red
# -*- acsection: after-main -*-
print(sati, minuta)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for sati, minuta, minuta_od_ponoci in [(14, 19, 859), (11, 13, 673), (23, 59, 1439)]:
          self.assertEqual(acMainSection(minuta_od_ponoci = minuta_od_ponoci)["sati"],sati,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
          self.assertEqual(acMainSection(minuta_od_ponoci = minuta_od_ponoci)["minuta"],minuta,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id8">
<h2>Аутобус<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Aутобус је кренуо са станице у 6 часова и 17 минута и путовао је 2
сата и 55 минута. У колико сати и минута је стигао?</p>

    </p></div></div>

<div data-childcomponent="време_путовања" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=време_путовања data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
# sat i munut kada je autobus krenuo
krenuo_sat = 6
krenuo_min = 17
# broj sati i minuta koliko je autobus putovao
putovao_sat = 2
putovao_min = 55
# sabiramo minute i prenosimo sate ako je potrebno (ako su minuti bar 60)
stigao_min = (krenuo_min + putovao_min) % 60;
prenos = (krenuo_min + putovao_min) // 60;
# sabiramo sate dodajuci prenos minuta
stigao_sat = krenuo_sat + putovao_sat + prenos;
print(stigao_sat, stigao_min)
</textarea>
</div>
<p>Други начин да се реши задатак је да се прво све прерачуна у минуте,
да се сабирање изврши у минутима, а да се затим добијени минути
претворе назад у сате и минуте. Приликом превођења неког временског
тренутка у минуте израчунавамо заправо број минута протеклих од
претходне поноћи (или поднева, ако се сати рачунају само до 12). Тај
број минута се може добити тако што се број сати помножи бројем 60
(јер у једном сату има 60 минута) и на то дода број минута.</p>

<div data-childcomponent="време_путовања_1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=време_путовања_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
# sat i minut kada je autobus krenuo
krenuo_sat = 6
krenuo_min = 17
# broj sati i minuta koliko je autobus putovao
putovao_sat = 2
putovao_min = 55
# trenutak polaska u minutima
krenuo_u_minutima = krenuo_sat * 60 + krenuo_min
# trajanje putovanja u minutima
putovao_u_minutima = putovao_sat * 60 + putovao_min
# trenutak dolaska u minutima
stigao_u_minutima = krenuo_u_minutima + putovao_u_minutima
# sat i minut dolaska
stigao_sat = stigao_u_minutima // 60
stigao_min = stigao_u_minutima % 60
print(stigao_sat, stigao_min)
</textarea>
</div>
</div>
<div class="section" id="id11">
<h2>Криви торањ у Пизи<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Криви торањ у Пизи је нагнут и са земљом заклапа угао од 86 степени
и 3 минута. Колико степени и минута је торањ нагнут, то јест,
колико одступа од усправног положаја.</p>

    </p></div></div>
<p>Прав угао има 90 степени тј. 90·60 минута. Ако угао од 86 степени и 3
минута преведемо у минуте и одузмемо га од правог угла добићемо
тражени угао у минутима. На крају ћемо израчунати целобројни количник
и остатак и тако добити тражени угао у степенима и минутама. Рецимо
још да смо у овом задатку заправо одређивали комплемент датог угла,
што је операција о којој је сигурно било речи на часовима математике.
Пошто се торањ с временом криви, твој програм треба исправно да ради и
ако је угао мало другачји.</p>

<div data-childcomponent="криви_торањ" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=криви_торањ data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="alfa_stepeni alfa_minuti beta_stepeni beta_minuti"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
# ugao pod kojim je nagnut toranj u stepenima i minutima
alfa_stepeni = 86
alfa_minuti = 3
# -*- acsection: main -*-
# prevodimo taj ugao u minute
alfa_u_minutima = 0      # ispravi ovaj red
# prav ugao u minutima
prav_u_minutima = 90 * 60
# komplement ugla u minutima
beta_u_minutima = prav_u_minutima - alfa_u_minutima
# komplement ugla u stepenima i minutima
beta_stepeni = 0         # ispravi ovaj red
beta_minuti = 0          # ispravi ovaj red
# -*- acsection: after-main -*-
print(beta_stepeni, "stepenа i", beta_minuti, "minuta")
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for (alfa_stepeni, alfa_minuti, beta_stepeni, beta_minuti) in [(86, 37, 3, 23), (85, 19, 4, 41)]:
          self.assertEqual(acMainSection(alfa_stepeni = alfa_stepeni, alfa_minuti = alfa_minuti)["beta_stepeni"],beta_stepeni,"Комплемент угла %s:%s је %s:%s." % (alfa_stepeni, alfa_minuti, beta_stepeni, beta_minuti))
          self.assertEqual(acMainSection(alfa_stepeni = alfa_stepeni, alfa_minuti = alfa_minuti)["beta_minuti"],beta_minuti,"Комплемент угла %s:%s је %s:%s." % (alfa_stepeni, alfa_minuti, beta_stepeni, beta_minuti))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id13">
<h2>Домаћи задатак<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h2>
<p>Ако на часу нисте стигли да урадите све задатке, уради их сада, у
склопу домаћег задатка. Након тога покушај да урадиш и наредних
неколико задатака.</p>
</div>
<div class="section" id="id14">
<h2>Прочитане стране књиге<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Књига има 282 стране. Марко је првог дана прочитао трећину, другог
дана половину остатка, а трећег књигу прочитао до краја. Колико
страна је прочитао ког дана? Напиши програм тако да исправно ради
и ако је број страна првог дана другачији.</p>

    </p></div></div>

<div data-childcomponent="Читање" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Читање data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="broj_strana prvi_dan drugi_dan treci_dan"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
broj_strana = 282
# -*- acsection: main -*-
prvi_dan = 0      # ispravi ovaj red
drugi_dan = 0     # ispravi ovaj red
treci_dan = 0     # ispravi ovaj red
# -*- acsection: after-main -*-
print(prvi_dan, drugi_dan, treci_dan)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for broj_strana, dan in [(369, 123), (333, 111)]:
          self.assertEqual(acMainSection(broj_strana = broj_strana)["prvi_dan"],dan,"Ако књига има %s страна, први дан је прочитано %s страна." % (broj_strana,dan))
          self.assertEqual(acMainSection(broj_strana = broj_strana)["drugi_dan"],dan,"Ако књига има %s страна, други дан је прочитано %s страна." % (broj_strana,dan))
          self.assertEqual(acMainSection(broj_strana = broj_strana)["treci_dan"],dan,"Ако књига има %s страна, трећи дан је прочитано %s страна." % (broj_strana,dan))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="пресек_решење11"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    <p>Марко је прочитао 94 стране сваког дана. Првог дана је прочитао
трећину и остале су му две трећине. Другог дана је прочитао пола од
тога тј. опет трећину и за трећи дан му је остала последња трећина.</p>

<div data-childcomponent="Читање_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Читање_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
broj_strana = 282
prvi_dan = broj_strana / 3
drugi_dan = (broj_strana - prvi_dan) / 2
treci_dan = broj_strana - prvi_dan - drugi_dan
print(treci_dan)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id17">
<h2>Повећање и смањење плата<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>
<p>Ево опет једног сличног задатка за вежбу.</p>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Плате су прво смањене за десет процената, а онда су после неколико
месеци те смањене плате повећане за 10 процената. Ако је у почетку
плата била 50,000 динара, колика је она после смањења и повећања?</p>

    </p></div></div>

<div data-childcomponent="плате" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=плате data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="plata povecana_plata"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
plata = 50000
# -*- acsection: main -*-
smanjena_plata = 0  # popravi ovaj red
povecana_plata = 0  # popravi ovaj red
# -*- acsection: after-main -*-
print(povecana_plata)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for plata, povecana_plata in [(45123, 44671.77), (54321, 53777.79)]:
          self.assertEqual(acMainSection(plata = plata)["povecana_plata"],povecana_plata,"Ако је плата била %s динара, после повећања она износи %s динара." % (plata,povecana_plata))
myTests().main()
</textarea>
</div>
<p>Ако добијеш решење 49,500 значи да је све како треба, иако је тај
резултат мало неочекиван. Често се помисли да ће смањење и повећање за
по 10% да се пониште и да ће се плата вратити на полазну
вредност. Међутим, смањење је било 10% полазне величине од 50,000
динара тј. за 5,000 динара, док је повећање за 10% од смањене величине
од 45,000 динара тј. повећање је за 450 динара.</p>

    <div data-component="reveal" id="пресек_решење21"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="плате_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=плате_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
plata = 50000
smanjena_plata = plata - (10 / 100) * plata
povecana_plata = smanjena_plata + (10 / 100) * smanjena_plata
print(povecana_plata)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id20">
<h2>Преостало време видео-снимка<a class="headerlink" href="#id20" title="Permalink to this headline">¶</a></h2>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Апликација за пуштање видео-снимака нуди опцију приказа преосталог
времена до краја снимка. Ако је познато трајање снимка у сатима,
минутима и секундама и време протекло од почетка снимка, такође
задато у сатима, минутима и секундама напиши програм који израчунава
време до краја у сатима, минутима и секундама.</p>

    </p></div></div>
<p>Најлакши начин да решиш задатак је да оба позната времена претвориш у
секунде, затим да тражено време израчунаш у секундама, а онда да
добијени резултат претвориш у сате, минуте и секунде.</p>

<div data-childcomponent="време_до_краја_видеа" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=време_до_краја_видеа data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="trajanje_sati trajanje_minuta trajanje_sekundi od_pocetka_sati od_pocetka_minuta od_pocetka_sekundi preostalo_sati preostalo_minuta preostalo_sekundi"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
trajanje_sati = 1
trajanje_minuta = 23
trajanje_sekundi = 14

od_pocetka_sati = 0
od_pocetka_minuta = 47
od_pocetka_sekundi = 53

# -*- acsection: main -*-
# izracunaj trajanje u sekundama
trajanje = 0

# izracunaj vreme od pocetka videa u sekundama
od_pocetka = 0

# izracunaj preostalo vreme u sekundama
preostalo = 0
# preracunaj to vreme u sate, minute i sekunde
preostalo_sati = 0
preostalo_minuta = 0
preostalo_sekundi = 0

# -*- acsection: after-main -*-
print(preostalo_sati, ":", preostalo_minuta, ":", preostalo_sekundi)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for (trajanje_sati, trajanje_minuta, trajanje_sekundi, od_pocetka_sati, od_pocetka_minuta, od_pocetka_sekundi, preostalo_sati, preostalo_minuta, preostalo_sekundi) in [(2, 54, 15, 1, 48, 29, 1, 5, 46), (1, 44, 13, 0, 0, 19, 1, 43, 54)]:
          self.assertEqual(acMainSection(trajanje_sati = trajanje_sati, trajanje_minuta = trajanje_minuta, trajanje_sekundi = trajanje_sekundi, od_pocetka_sati = od_pocetka_sati, od_pocetka_minuta = od_pocetka_minuta, od_pocetka_sekundi = od_pocetka_sekundi)["preostalo_sati"],preostalo_sati,"Ако је трајање (%s, %s, %s), а од почетка је протекло (%s, %s, %s), тада је преостало (%s, %s, %s)." % (trajanje_sati, trajanje_minuta, trajanje_sekundi, od_pocetka_sati, od_pocetka_minuta, od_pocetka_sekundi, preostalo_sati, preostalo_minuta, preostalo_sekundi))
          self.assertEqual(acMainSection(trajanje_sati = trajanje_sati, trajanje_minuta = trajanje_minuta, trajanje_sekundi = trajanje_sekundi, od_pocetka_sati = od_pocetka_sati, od_pocetka_minuta = od_pocetka_minuta, od_pocetka_sekundi = od_pocetka_sekundi)["preostalo_minuta"],preostalo_minuta,"Ако је трајање (%s, %s, %s), а од почетка је протекло (%s, %s, %s), тада је преостало (%s, %s, %s)." % (trajanje_sati, trajanje_minuta, trajanje_sekundi, od_pocetka_sati, od_pocetka_minuta, od_pocetka_sekundi, preostalo_sati, preostalo_minuta, preostalo_sekundi))
          self.assertEqual(acMainSection(trajanje_sati = trajanje_sati, trajanje_minuta = trajanje_minuta, trajanje_sekundi = trajanje_sekundi, od_pocetka_sati = od_pocetka_sati, od_pocetka_minuta = od_pocetka_minuta, od_pocetka_sekundi = od_pocetka_sekundi)["preostalo_sekundi"],preostalo_sekundi,"Ако је трајање (%s, %s, %s), а од почетка је протекло (%s, %s, %s), тада је преостало (%s, %s, %s)." % (trajanje_sati, trajanje_minuta, trajanje_sekundi, od_pocetka_sati, od_pocetka_minuta, od_pocetka_sekundi, preostalo_sati, preostalo_minuta, preostalo_sekundi))
myTests().main()
</textarea>
</div>
<p>Ако урадиш све како треба, добићеш да је преостало <code class="docutils literal notranslate"><span class="pre">0</span> <span class="pre">:</span> <span class="pre">35</span> <span class="pre">:</span>
<span class="pre">21</span></code>. Провери и на другим тест-примерима.</p>

    <div data-component="reveal" id="пресек_решење41"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="време_до_краја_видеа_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=време_до_краја_видеа_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
# ukupno vreme trajanja videa
trajanje_sati = 1
trajanje_minuta = 23
trajanje_sekundi = 14

# vreme proteklo od pocetka videa
od_pocetka_sati = 0
od_pocetka_minuta = 47
od_pocetka_sekundi = 53

# izracunaj trajanje u sekundama
trajanje = (trajanje_sati*60 + trajanje_minuta)*60 + trajanje_sekundi

# izracunaj vreme od_pocetka od pocetka videa u sekundama
od_pocetka = (od_pocetka_sati*60 + od_pocetka_minuta)*60 + od_pocetka_sekundi

# izracunaj preostalo vreme do kraja videa u sekundama
preostalo = trajanje - od_pocetka
# preracunaj to vreme u sate, minute i sekunde
preostalo_sati = preostalo // (60 * 60)
preostalo_minuta = (preostalo // 60) % 60
preostalo_sekundi = preostalo % 60

print(preostalo_sati, ":", preostalo_minuta, ":", preostalo_sekundi)
</textarea>
</div>

    </div>
    </div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - 9.1. Цели и реални бројеви, дељење' data-toggle="tooltip"><a href="9a.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - 10. лекција' data-toggle="tooltip"><a href="../10_lekcija/index.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - 9.1. Цели и реални бројеви, дељење' data-toggle="tooltip"><a href="9a.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - 10. лекција' data-toggle="tooltip"><a href="../10_lekcija/index.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>